Fat body-derived cytokine Upd2 controls disciplined migration of tracheal stem cells in Drosophila

Author response:

The following is the authors’ response to the previous reviews

Joint Public Review:

Summary

In this manuscript, Dong et al. study the directed cell migration of tracheal stem cells in Drosophila pupae. The authors study how the directionality of these cells is regulated along the dorsal trunk. They show that inter-organ communication between the tracheal stem cells and the nearby fat body plays a role in posterior migration. They provide compelling evidence that Upd2 production in the fat body and JAK/STAT activation in the tracheal stem cells play a role. Moreover, they show that JAK/STAT signalling might induce the expression of apicobasal and planar cell polarity genes in the tracheal stem cells which appear to be needed to ensure unidirectional migration. Finally, the authors suggest that trafficking and vesicular transport of Upd2 from the fat body towards the tracheal cells might be important.

Strengths

The manuscript is well written and presents extensive and varied experimental data to show a link between Upd2-JAK/STAT signaling from the fat body and tracheal progenitor cell migration. The authors provide convincing evidence that the fat body, located near the trachea, secretes vesicles containing the Upd2 cytokine and that affecting JAK-STAT signaling results in aberrant migration of some of the tracheal stem cells towards the anterior. Using ChIP-seq as well as analysis of GFP-protein trap lines of planar cell polarity genes in combination with RNAi experiments, the authors show that STAT92E likely regulates the transcription of planar cell polarity genes and some apicobasal cell polarity genes in tracheal stem cells which appear to be needed for unidirectional migration. The work presented here provides some novel insights into the mechanism that ensures polarized migration of tracheal stem cells, preventing bidirectional migration. This might have important implications for other types of directed cell migration in invertebrates or vertebrates including cancer cell migration. Overall, the authors have substantially improved their manuscript since the first submission but there are still some weaknesses.

Weaknesses

Overall, the manuscript lacks insights into the potential significance of the observed phenotypes and of the proposed new signaling model. Most of our concerns could be dealt with by adjusting the text (explaining some parts better and toning down some statements).

(1) Directional migration of tracheal progenitors is only partially compromised, with some cells migrating anteriorly and others maintaining their posterior migration, a quite discrete phenotype.

The strongest migration defects quantified in graphs (e.g. 100 μm) are not shown in images, since they would be out of frame, it would be beneficial to see them. In addition, the consequence of defects in polarized migration on tracheal development is not clear and data showing phenotypes on the final trachea morphology in pupae are not explained nor linked to the previous phenotypes.

We agree with you that it is informative to show strong anterior migration (> 100 μm). Accordingly, we have shown examples in Figure 3B and Figure 7R-S. In addition, we have also discuss on the links between migration defects and the consequential phenotypes of the animal at a later developmental stage in the revised manuscript. The undisciplined migration leads to insufficient regeneration and incomplete remodeling of airway and causes pupal lethality.

(2) Some important information is lacking, such as the origin of mutant and UAS-RNAi lines, which are not reported in the material and methods. For instance, mutants for components of the JAK-STAT pathway are used but not described. Are they all viable at the pupal stage? Otherwise, pupae would not be homozygous mutants. From the figure legend, it seems that the Stat92EF allele has been used, which is a point mutation, thus not leading to an absence of protein. If the hopTUM allele has been used, as mentioned in the legend, it is a gain-of-function allele. Thus, the authors should not conclude that "The aberrant anterior migration of tracheal progenitors in the absence of JAK/STAT components led to impairment of tracheal integrity and caused melanization in the trachea (Figure 3-figure supplement 1E-I)".

We apologize for inadequate description of the experimental materials and methods. We have listed the stock number of mutant and RNAi alleles in Key resource table and Materials. The mutant alleles that we chose to examine can survive to pupal stage, which is key to the success of our subsequent characterization of these mutants. According to your suggestion, we modified the statement for accuracy.

(3) The authors observe that tracheal progenitors display a polarized distribution of Fat that is controlled by JAK-STAT signaling. However, this conclusion is made from a single experiment using only 3 individuals with no statistics. This is insufficient to support the claim that "JAK/STAT signaling promotes the expression of genes involved in planar cell polarity leading to asymmetric localization of Fat in progenitor cells", as mentioned in the abstract, or that "the activated tracheal progenitors establish a disciplined migration through the asymmetrical distribution of polarity proteins which is directed by an Upd2-JAK/STAT signaling stemming from the remote organ of fat body."

We performed multiple biological replicates for Ft distribution experiments and observed similar trend, although we only showed three representative samples. In the revised text, we have included n number for statistic representation and statistic test.

(4) The authors demonstrate that Upd2 is transported through vesicles from the fat body to the tracheal progenitors. It remains somewhat unclear in the proposed model how Upd2 activates JAK-STAT signaling. Are vesicles internalized, as it seems to be proposed, and thus how does Upd2 activate JAK-STAT signaling intracellularly? Or is Upd2 released from vesicles to bind Dome extracellularly to activate the JAK-STAT pathway? Moreover, it is not clear nor discussed what would be the advantage of transporting the ligand in vesicles compared to classical ligand diffusion.

We do not know whether the association between Upd2 and Lbm is inside or outside vesicles. The vesicular trafficking of Upd2 is our observation and supported by various genetic and biochemical experiments. Our research does not imply the message that this vesicular trafficking has advantage over diffusion.

Changes from the earlier version to the revised version

Abstract

Abstract

Coordinated activation and directional migration of adult stem cells are essential for maintaining tissue homeostasis. Drosophila tracheal progenitors are adult stem cells that migrate posteriorly along the dorsal trunk to replenish degenerating branches that disperse the FGF mitogen. However, it is currently unknown how the overall anterior-to-posterior directionality of such migration is controlled. Here we show that individual progenitor cells migrate together in a concerted, disciplined manner, a behavior that is dependent on the neighboring fat body. We identify the fat body-derived cytokine, Upd2, in targeting and inducing JAK/STAT signaling in tracheal progenitors to maintain their directional migration. Perturbation of either Upd2 production in fat body or JAK/STAT signaling in trachea causes aberrant bidirectional migration of tracheal progenitors. We show that JAK/STAT signaling promotes the expression of genes involved in planar cell polarity leading to asymmetric localization of Fat in progenitor cells. We provide evidence that Upd2 transport requires Rab5- and Rab7-mediated endocytic sorting and Lbm-dependent vesicle trafficking. Our study thus uncovers an inter-organ communication in the control of disciplined migration of tracheal progenitor cells, a process that requires vesicular trafficking of fat body-derived cytokine Upd2 and JAK/STAT signaling-mediated activation of PCP genes.

Introduction

Adult stem cells are multipotent cell populations which inhabit their niche but mobilize to initiate tissue reconstruction during organismal growth and regeneration. An intriguing feature of stem cells is their capability of migrating in a disciplined directionality towards locations undergoing reconstruction (Li and Clevers, 2010). Such a highly disciplined movement is critical for maintaining tissue homeostasis and is influenced by various niche-intrinsic signals and external stimuli, and its aberrancy causes diseases such as hypertrophy (Zhou et al., 2024). The damaged tissue or distant organs that elicit systemic signals promote the migration of adult stem cells (Jones and Wagers, 2008). In addition, interactions with other cell types, soluble factors (e.g. cytokines, growth factors, and hormones) and tissue stiffness collectively bolster the mobilization of stem cells (Fuchs and Blau, 2020). Despite growing appreciation of adult stem cells as a primary source for tissue regeneration, the mechanism governing directional stem cell migration remains yet to be elucidated.

Earlier version: Drosophila tracheal progenitors are a population of adult stem cells that rebuild the degenerating trachea during metamorphosis. The progenitor cells reside in Tr4 and Tr5 metameres and start to move along the tracheal branch towards sites of regeneration (Chen and Krasnow, 2014; Pitsouli and Perrimon, 2010). Movement of these progenitor cells follows a stereotypical anterior-to-posterior axis (Figure 1A), thus representing a suitable system to investigate mechanisms controlling the directionality of stem cell migration. The activation of tracheal progenitors is stimulated by the morphogen Branchless (Bnl), fly homolog of fibroblast growth factor (FGF) (Chen and Krasnow, 2014), and the insulin hormone (Li et al., 2022). Intercellular communication and synergy between organs also contribute to the branching morphogenesis (Perochon et al., 2021; Schottenfeld et al., 2010; Tamamouna et al., 2021). The functional role of the interactions between trachea and other organs in modulating tracheal progenitor behavior has been largely unknown. Drosophila fat body is the functional analog of mammalian adipose tissue and the major organ sensing various hormonal and nutritional signals to orchestrate systemic growth, metabolism and stem cell maintenance (Sriskanthadevan-Pirahas et al., 2022). Fat body produces regulatory molecules known as fat body signals (FBSs), which remotely affect the activity of other organs (Ingaramo et al., 2020; Zheng et al., 2016). For instance, the fat body-to-brain signals modulate insulin-like peptides production (Rajan and Perrimon, 2012), visual attention and sleep behavior (Ertekin et al., 2020). The Drosophila family of interleukin-6 (IL-6)-like cytokines consist of Unpaired (Upd, also called Outstretched), Upd2 and Upd3, and serve as mediators of systemic signaling. Whereas Upd1 and Upd3 derive from fly brain and plasmatocytes (Beshel et al., 2017; Woodcock et al., 2015), Upd2 is primarily produced by the fat body (Rajan et al., 2017), although muscle-derived Upd2 is also reported (Zhao and Karpac, 2017). The Upd proteins act as ligands which bind to a common GP130-like receptor, Domeless (Dome) on target cells (Agaisse et al., 2003; Chen et al., 2002). Upon association of ligands, the Dome receptors dimerize and recruit the non-receptor tyrosine kinase JAKs leading to their subsequent transactivation via phosphorylation. The transactivated JAKs then phosphorylate the tyrosine residues of their substrates, including the bound receptors and cytosolic STATs. The phosphorylation of STATs promotes their dimerization and nuclear translocation to activate transcriptional program (Darnell, 1997). JAK/STAT signaling requires the IL-6 cytokines (Heinrich et al., 2003), and is implicated in numerous cellular events including cell proliferation, differentiation, migration and apoptosis (O’Shea et al., 2002).

Revised version: Drosophila tracheal progenitors are a population of adult stem cells that rebuild the degenerating trachea during metamorphosis. The progenitor cells reside in Tr4 and Tr5 metameres and start to move along the tracheal branch towards sites of regeneration (Chen and Krasnow, 2014; Pitsouli and Perrimon, 2010). Movement of these progenitor cells follows a stereotypical anterior-to-posterior axis (Figure 1A), thus representing a suitable system to investigate mechanisms controlling the directionality of stem cell migration. The activation of tracheal progenitors is stimulated by the morphogen Branchless (Bnl), fly homolog of fibroblast growth factor (FGF) (Chen and Krasnow, 2014), and the insulin hormone (Li et al., 2022). Intercellular communication and synergy between organs also contribute to the branching morphogenesis (Perochon et al., 2021; Schottenfeld et al., 2010; Tamamouna et al., 2021). The functional role of the interactions between trachea and other organs in modulating tracheal progenitor behavior has been largely unknown. Drosophila fat body is the functional analog of mammalian adipose tissue and the major organ sensing various hormonal and nutritional signals to orchestrate systemic growth, metabolism and stem cell maintenance (Sriskanthadevan-Pirahas et al., 2022).

Fat body affects disciplined migration of tracheal progenitors.

(A) Schematic cartoon showing the migration of tracheal progenitors (red) and degenerative tracheal branches (dashed grey lines) in pupae. Fat body is shown in beige. Arrows denote anterior-posterior (A-P) axis. (B, C) Frontal section (B) and sagittal view (C) showing the relative position of fat body and tracheal progenitors. (D-J’) Migration of tracheal progenitors in control and fat body perturbation flies. (D-D’’) Migration of tracheal progenitors (red) upward from transverse connective (blue dashed lines) and along the dorsal trunk (white dotted lines) at 0hr APF (D), 1hr APF (D’), and 3hr APF (D’’). (E-E’’) Bidirectional movement of tracheal progenitors in fat body-depleted (lsp2>rpr.hid) flies. 0hr APF (E), 1hr APF (E’), 3hr APF (E’’). Arrows point to anterior movement of tracheal progenitors. (F) Bar graph showing the migration distance of tracheal progenitors. The top chart of column represents the migration distance of anterior-most stem cells, and the lower chart of column represents the migration distance of posterior-most stem cells. Error bars represent SEM, n=6. (G-G’) The distribution of progenitors at 2hr APF. (H-H’’) The distribution of progenitors in fat body-depleted flies at 2hr APF. (I-J’) Computer simulation depicting trajectories of progenitor migration. (I, J) Confocal images of tracheal progenitors. (I’, J’) Vectors of progenitor migration. (K) Bar graph plots the binary entropy that represents the disorderedness of migration direction of tracheal progenitors. Error bars represent SEM, n=10. (L) The Bernoulli random variable X showing optic flow distribution of the binarized directions in each group. Error bars represent SEM, n=10. N.S. indicates not significant. Scale bar: 100 μm (B, C, G), 200 μm (D-E’’). Genotypes: (B, C) UAS-mCD8-GFP/+; lsp2-Gal4,P[B123]-RFP-moe/+; (D-D’’, G, G’) Gal80ts/+;lsp2-Gal4,P[B123]-RFP-moe/+ ; (E-E’’, H, H’) UAS-rpr-hid/+;Gal80ts/+;lsp2-Gal4,P[B123]-RFP-moe/+.

Added in the revised version: Fat body produces regulatory molecules known as fat body signals (FBSs), which remotely affect the activity of other organs (Ingaramo et al., 2020; Zheng et al., 2016). For instance, the fat body-to-brain signals modulate insulin-like peptides production (Rajan and Perrimon, 2012), visual attention and sleep behavior (Ertekin et al., 2020). The Drosophila family of interleukin-6 (IL-6)-like cytokines consist of Unpaired (Upd, also called Outstretched), Upd2 and Upd3, and serve as mediators of systemic signaling. Whereas Upd1 and Upd3 derive from fly brain and plasmatocytes (Beshel et al., 2017; Woodcock et al., 2015), Upd2 is primarily produced by the fat body (Rajan et al., 2017), although muscle-derived Upd2 is also reported (Zhao and Karpac, 2017). The Upd proteins act as ligands which bind to a common GP130-like receptor, Domeless (Dome) on target cells (Agaisse et al., 2003; Chen et al., 2002). Upon association of ligands, the Dome receptors dimerize and recruit the non-receptor tyrosine kinase JAKs leading to their subsequent transactivation via phosphorylation. The transactivated JAKs then phosphorylate the tyrosine residues of their substrates, including the bound receptors and cytosolic STATs. The phosphorylation of STATs promotes their dimerization and nuclear translocation to activate transcriptional program (Darnell, 1997). JAK/STAT signaling requires the IL-6 cytokines (Heinrich et al., 2003), and is implicated in numerous cellular events including cell proliferation, differentiation, migration and apoptosis (O’Shea et al., 2002).

Here, we investigate molecular basis underlying directional stem cell migration using the Drosophila tracheal progenitors as a model. Our results identify a cytokine-mediated inter-organ communication between fat body and the progenitor cells that is necessary for their disciplined, directional migration. The directional migration of the progenitors relies on JAK/STAT signaling and its downstream targets of planar cell polarity (PCP) components. Importantly, the Upd2 cytokines derived from fat body are transported through vesicular trafficking to induce JAK/STAT signaling in tracheal progenitors. Our study reveals that tracheal progenitors establish migratory directionality as they exit their niches and that the disciplined migration of the progenitors depends on an inter-organ signaling originating from the fat body.

Results

Dependence of tracheal progenitors on the fat body

The fly tracheal progenitors are activated and move posteriorly along the dorsal trunk (DT) at the onset of pupariation (Figure 1A). We set out to delve into the underlying mechanisms of directional progenitor cell movement and tentatively surveyed organs that may coordinate this process. In Drosophila, the fat body resides anatomically in proximity with trachea (Figure 1B, C; Video 1) and is the principal reservoir for energy consumption. To determine whether the integrity of fat body is required for tracheal progenitors, we perturbed larval or pupal fat body by expressing pro-apoptotic cell death genes, hid and reaper (rpr), under the control of a fat body-specific driver, lsp2-Gal4 (Cherbas et al., 2003). Expression of hid and rpr in L3 stage impaired fat body integrity and adipocyte abundance, and generated slender larvae and pupae (Figure 1-figure supplement 1A-D). In these animals, the tracheal progenitors exhibited a sign of undisciplined migration and tended to move bidirectionally (Figure 1D-F), although their migration rate, cell number and proliferation remained unchanged (Figure 1-figure supplement 1E-H and Figure 3-figure supplement 1). The undisciplined bidirectional migration behavior of tracheal progenitors in fat body-defective animals is in stark contrast to control animals where the progenitors migrated unambiguously towards posterior (Figure 1D-D’’ and Video 2). To gain a quantitative view of progenitor cell migration, we traced the movement of individual cells by time-lapse confocal imaging. At 2hr APF, tracheal progenitors from fat body deficit animals displayed a symmetrical distribution relative to the junction between dorsal trunk (DT) and transverse connective (TC), compared with an L-shape localization of niche-associated and migratory progenitors established by a posterior movement in control (Figure 1G-H’). Gauging the vector denoting the movement of each progenitor (Figure 1I-I’) revealed that the traces of progenitor groups in fat body-depleted animals exhibited a fan-shaped pattern (Figure 1J, J’). Owing to this undisciplined movement, entropy of the system was notably elevated upon increased inconsistency of migration vectors (Figure 1K). The bidirectional migratory progenitors displayed longer territory as assessed by binarized direction (Figure 1L). Collectively, these results suggest that fat body has an integral role in maintaining the discipline of tracheal progenitor movement.

Upd2-JAK/STAT signaling between fat body and trachea

Since fat body impacts the behavior of tracheal progenitors, we next attempted to investigate the signal between these two interdependent organs. For this purpose, we first performed RNA sequencing (RNA-seq) analysis of tracheal progenitors from aforementioned fat body defective flies. The results revealed a dramatical alteration of transcriptional program in tracheal progenitors upon the perturbation of fat body (Figure 2-figure supplement 1A, B). Interestingly, the functional cluster of ‘cytokine activity’ showed prominent enrichment in the differentially expressed genes (DEGs) in progenitors from lsp2>rpr.hid pupae (Figure 2A). This raised the possibility that certain cytokine-responsive signaling was induced in tracheal progenitors and the signaling was compromised by impairment of fat body. Therefore, we proceeded to analyze the expression of genes responsive to cytokine signaling. Analyzing the RNA-seq data revealed that the cytokine-dependent JAK/STAT and Dpp signaling were notably up-regulated upon the activation of progenitors (Figure 2B). Importantly, fat body depletion led to suppression of target genes of JAK/STAT, PI3K and Dpp signaling in tracheal progenitors, suggesting their dependence on the function of fat body (Figure 2C).

Dependence of tracheal progenitors on cytokines from fat body.

(A) Top functional clusters among the differentially expressed genes of progenitors between control and fat body-depleted pupae. Gene ratio refers to the proportion of genes in a dataset that are associated with a particular biological process, function, or pathway. Count indicates the number of genes from an input gene list that are associated with a specific GO term. (B) Heatmap depicting expression levels of principal target genes of signaling pathways in L3 larvae, 0hr APF pupae and 2hr APF pupae. (C) Heatmap showing the differential expression of target genes of signaling pathways between control and fat body-depleted pupae. (D-D’’) Migration of tracheal progenitors along the dorsal trunk at 0hr APF (D), 1hr APF (D’), and 3hr APF (D’’). The white dashed line shows transverse connective. (E-E’’). Migration of tracheal progenitors in upd2RNAi flies. (F) Bar graph plots the migration distance of tracheal progenitors. Error bars represent SEM, n=6. (G) Volcano plot showing surface proteomics of tracheal epithelium (up-regulated genes with ten-fold or higher changes in red; down-regulated genes with ten-fold or higher changes in blue). (H) Top functional classes among the surface proteomics of trachea. (I) Schematic diagram depicting the working principle of the DIPF reporter. (J) The signal of DIPF reporter in tracheal progenitors. The progenitors are outlined by dashed lines. N.S. indicates not significant. Scale bar: 200 μm (D-E’’), 50 μm (J). Genotypes: (A, C) lsp2-Gal4,P[B123]-RFP-moe/+for control, UAS-rpr-hid/+;Gal80ts/+;lsp2-Gal4,P[B123]-RFP-moe/+; (B) P[B123]-RFP-moe/+. (D-D’’) lsp2-Gal4,P[B123]-RFP-moe/+; (E-E’’) lsp2-Gal4,P[B123]-RFP-moe/UAS-upd2RNAi; (J) btl-Gal4/UAS-DIPF.

To evaluate the roles of these signaling proteins, we perturbed their expression in fat body by the expression of RNAi constructs. Knockdown of candidates including some cytokines specifically in fat body did not affect the direction of tracheal progenitor migration (Figure 2-figure supplement 1C-J), except for upd2, whose depletion phenocopied fat body ablation-induced bidirectional movement of tracheal progenitors (Figure 2D-F and Video 3). These results suggest a role of fat body-produced Upd2 in remotely regulating the tracheal progenitors.

Then, we performed surface proteome in vivo (Li et al., 2020) to investigate the spectrum of molecules received by trachea (Figure 2-figure supplement 2). The trachea-associated proteins were biotinylated through a reaction mediated by a membrane-tethered horse radish peroxidase (HRP-CD2) (Figure 2-figure supplement 2A-D). Of the 1684 streptavidin-precipitated proteins captured by mass spectrometry (Figure 2G), a functional cluster enriched for receptor signaling via JAK/STAT was identified (Figure 2H). The JAK/STAT pathway is one of the principal cellular signaling that responds to Upd2 ligand (Hombria et al., 2005). Drosophila JAK/STAT signaling is well-conserved (Arbouzova and Zeidler, 2006; Zeidler et al., 2000) and comprises a single JAK (Hopscotch, Hop) and one STAT (Stat92E), in contrast to a handful of homologues (four JAK and seven STAT genes) found in mammals. Domeless (Dome), the receptor for JAK/STAT pathway, exhibited pronounced expression in the tracheal progenitors (Figure 2-figure supplement 2E). To test if these Dome receptors actively interact with their ligands, we adapted a technique to monitor ligand-receptor interaction in vivo (Michel et al., 2011) and constructed a Dome variant (DIPF) which only fluoresces in the ligand-binding and phosphorylated state (Figure 2I). The signal of this DIPF reporter was detected in both larval fat body and salivary gland (Figure 2-figure supplement 2F,H), which is consistent with active JAK signaling implicated in the development of the tissues (Chakrabarti et al., 2016; Krautz et al., 2020). When expressed in the tracheal system, DIPF displayed robust fluorescent signal in the tracheal progenitors (Figure 2J). These data suggest that receptor signaling of JAK/STAT is active in the tracheal progenitors.

Earlier version: To analyze the functional importance of JAK/STAT signaling in tracheal progenitors, we perturbed the principal components of this signaling, namely the receptor Dome, signal transducer Hop or the downstream transcription factor Stat92E, by btl-Gal4-driven expression of RNAi constructs. Under these conditions in which JAK/STAT pathway is compromised, the tracheal progenitors aberrantly migrated anteriorly, which is reminiscent of upd2 loss-of-function in the fat body (Figure 3A-E, Video 4 and Figure 3-figure supplement1A-D). The aberrant anterior migration of tracheal progenitors in the absence of JAK/STAT components led to impairment of tracheal integrity and caused melanization in the trachea (Figure 3-figure supplement 1E-I). In agreement with genetic perturbation of JAK/STAT signaling, pharmacological inhibition of JAK by a small-molecule inhibitor, Tofacitinib (Palmroth et al., 2021), also triggered bidirectional movement of tracheal progenitors (Figure 3F-H). The bidirectional movement was not due to excessive progenitors or crowding (Figure 3-figure supplement 2A).

Revised version: To analyze the functional importance of JAK/STAT signaling in tracheal progenitors, we perturbed the principal components of this signaling, namely the receptor Dome, signal transducer Hop or the downstream transcription factor Stat92E, by btl-Gal4-driven expression of RNAi constructs. Under these conditions in which JAK/STAT pathway is compromised, the tracheal progenitors aberrantly migrated anteriorly, which is reminiscent of upd2 loss-of-function in the fat body (Figure 3A-E, Video 4 and Figure 3-figure supplement1A-D). The aberrant anterior migration of tracheal progenitors upon perturbation of JAK/STAT components led to incomplete regeneration of airway and impairment of tracheal integrity and caused melanization in the trachea (Figure 3-figure supplement 1E-I). In agreement with genetic perturbation of JAK/STAT signaling, pharmacological inhibition of JAK by a small-molecule inhibitor, Tofacitinib (Palmroth et al., 2021), also triggered bidirectional movement of tracheal progenitors (Figure 3F-H). The bidirectional movement was not due to excessive progenitors or crowding (Figure 3-figure supplement 2A).

JAK/STAT pathway is required for the discipline of tracheal progenitor migration.

(A-D’’) Migration of tracheal progenitors along the dorsal trunk at 0hr APF, 1hr APF and 3hr APF. The white dashed line shows transverse connective. The progenitors of control (A-A’’), domeRNAi (B-B’’), hopRNAi (C-C’’), and stat92ERNAi (D-D’’) flies. (E) Bar graph showing migration distance of progenitors. Error bars represent SEM, n=6. (F-G’’) JAK inhibition causes bidirectional movement of progenitors. Migration of tracheal progenitors in the absence (DMSO-fed) (F-F’’) or in the presence of Tofacinib (JAK inhibitor) (G-G’’). (H) Bar graph showing the distance of anterior movement. Error bars represent SEM, n=6. Scale bar: 200 μm (A-D’’, F-G’’). Genotypes: (A-A’’, F-G’’) btl-Gal4/+;P[B123]-RFP-moe/+; (B-B’’) btl-Gal4/+;P[B123]-RFP-moe/UAS-domeRNAi; (C-C’’) btl-Gal4/+;P[B123]-RFP-moe/UAS-hopRNAi; (D-D’’) btl-Gal4/+;P[B123]-RFP-moe/UAS-stat92ERNAi.

Concurrently, the activity of JAK/STAT pathway, as assessed by the Stat92E-GFP reporter (Bach et al., 2007), was substantially impaired when components of the pathway were depleted (Figure 3-figure supplement 2B-E, J). To determine whether the tracheal JAK/STAT signaling depends on fat body-derived Upd2, we depleted Upd2 in fat body and observed that Stat92E-GFP signal in tracheal progenitors was severely decreased, suggesting that JAK/STAT signaling in the trachea requires fat body-produced Upd2 (Figure 3-figure supplement 2F, G, J). Consistently, inhibition of JAK/STAT signaling using Tofacitinib reduced the expression of Stat92E-GFP (Figure 3-figure supplement 2H-J). Taken together, these observations suggest that Upd2-responsive JAK/STAT signaling in the trachea is essential for the disciplined migration of progenitors.

Genes regulated by JAK/STAT signaling in the trachea

Earlier version: To gain a comprehensive understanding of the molecular details underlying the discipline of tracheal progenitor migration, we conducted genomic chromatin immunoprecipitation (ChIP-seq) to identify loci bound by Stat92E which functions as the transcription factor of JAK/STAT pathway. This revealed a total of 21312 Stat92E binding peaks, ∼95.7% of which located within 2LJkb of transcription start sites (TSSs) of annotated genes (Figure 4-figure supplement 1A). In particular, 86% of the peaks (18328 peaks) were enriched either in promoter regions or within gene bodies, and 66.1% of the peaks (13490 peaks) resided near the 5LJ ends of annotated genes, namely in the promoter regions, first exons and first introns (Figure 4-figure supplement 1B, C). GO analysis of putative target genes of Stat92E identified one cluster associated with establishment of planar polarity (Figure 4A). In line with this, the functional class associated with establishment of planar polarity was also abundantly represented among the differentially expressed genes (DEGs) upon the activation of tracheal progenitors in larval-pupal transition (Figure 4B). Notably, Stat92E binding was detected in the promoters and intronic regions of genes functioning in distal-to-proximal signaling (Cho and Irvine, 2004), such as dachsous (ds), four-jointed (fj), fz, stan, Vang and fat2 (Figure 4C). Additionally, Stat92E occupied in the promoter regions of crb and yurt, two genes involved in apical-basal polarity and tracheal tube growth (Laprise et al., 2006; Schottenfeld-Roames and Ghabrial, 2012; Schottenfeld-Roames et al., 2014) (Figure 4-figure supplement 1D). The enrichment of Stat92E in the promoters and/or regulatory regions of these putative targets was confirmed by ChIP-qPCR (Figure 4-figure supplement 1E). To further validate these putative Stat92E targets and investigate their dependence on JAK/STAT signaling, we analyzed their expression from several fosmid transgenes which have a GFP tag fused to ds, fj or ft and express at endogenous levels. Ds and Fj were abundant in the progenitor cells, but were vastly reduced upon depletion of dome, hop or stat92E, suggesting that they are regulated by JAK/STAT pathway (Figure 4D-M). Furthermore, it is reported that the function of Ft is influenced by cell-autonomous increase of Ds level and its protein level is enhanced by Ds reduction (Ambegaonkar et al., 2012; Matakatsu and Blair, 2004), which is also evidenced by our analysis using dsRNAi and UAS-ds in the tracheal progenitors (Figure 4-figure supplement 2A-D). In accordance with this notion, the level of Ft, as assayed by the Ft-GFP reporter, was elevated by the reduction of JAK/STAT signaling (Figure 4N-R). We also analyzed GFP-tagged fosmid transgenes of fat2, crb and yurt and found that they were discernably reduced upon impairment of JAK/STAT signaling, suggesting that they are also regulated by JAK/STAT pathway (Figure 4-figure supplement 1F-T). Additionally, the transcription of ds, fj, ft, fat2, crb, and yurt was compromised by expression of stat92ERNAi (Figure 4-figure supplement 1U). In sum, these results suggest that JAK/STAT promotes components involved in the establishment of polarity in tracheal cells.

Revised version: To gain a comprehensive understanding of the molecular details underlying the discipline of tracheal progenitor migration, we conducted genomic chromatin immunoprecipitation (ChIP-seq) to identify loci bound by Stat92E which functions as the transcription factor of JAK/STAT pathway. This revealed a total of 21312 Stat92E binding peaks, ∼95.7% of which located within 2 kb of transcription start sites (TSSs) of annotated genes (Figure 4-figure supplement 1A). In particular, 86% of the peaks (18328 peaks) were enriched either in promoter regions or within gene bodies, and 66.1% of the peaks (13490 peaks) resided near the 5ʹ ends of annotated genes, namely in the promoter regions, first exons and first introns (Figure 4-figure supplement 1B, C). GO analysis of putative target genes of Stat92E identified one cluster associated with establishment of planar polarity (Figure 4A). In line with this, the functional class associated with establishment of planar polarity was also abundantly represented among the differentially expressed genes (DEGs) upon the activation of tracheal progenitors in larval-pupal transition (Figure 4B). Notably, Stat92E binding was detected in the promoters and intronic regions of genes functioning in distal-to-proximal signaling (Cho and Irvine, 2004), such as dachsous (ds), four-jointed (fj), fz, stan, Vang and fat2 (Figure 4C). Additionally, Stat92E occupied in the promoter regions of crb and yurt, two genes involved in apical-basal polarity and tracheal tube growth (Laprise et al., 2006; Schottenfeld-Roames and Ghabrial, 2012; Schottenfeld-Roames et al., 2014) (Figure 4-figure supplement 1D). The enrichment of Stat92E in the promoters and/or regulatory regions of these putative targets was confirmed by ChIP-qPCR (Figure 4-figure supplement 1E).

Identification of gene targets of Stat92E in Drosophila tracheal progenitors.

(A) Bubble plot represents the top functional clusters among gene targets. The establishment of planar polarity denoted in red solid box is identified with high enrichment score. (B) Top functional classes among the differentially expressed genes in larval-pupal transition. (C) ChIP-seq peaks at loci regulated by Stat92E. Scale bar: 20kb (ds, fz, stan), 5kb (fat2, vg), 1kb (fj). (D-H) Validation of gene targets of Stat92E ChIP-seq. (D-G) The expression of Ds-GFP in the tracheal progenitors of control (D), domeRNAi (E), hopRNAi (F) and stat92ERNAi (G). The progenitors are outlined by dashed lines. (H) The bar graphs plot the relative level of Ds. Error bars represent SEM, n=6. (I-M) The expression of Fj in tracheal progenitors. The expression of Fj-GFP in the tracheal progenitors of control (I), domeRNAi (J), hopRNAi (K) and stat92ERNAi (L). Dashed lines outline tracheal progenitors. (M) The bar graphs plot the relative level of Fj. Error bars represent SEM, n=6. (N-R) The level of Ft-GFP in the tracheal progenitors of control (N), domeRNAi (O), hopRNAi (P) and stat92ERNAi (Q). (R) The bar graphs plot the relative level of Ft. Error bars represent SEM, n=4. Scale bar: 50 μm (D-G, I-L, N-Q). Genotypes: (D) btl-Gal4,Ds-GFP/+; (E) btl-Gal4,Ds-GFP/+;UAS-domeRNAi/+; (F) btl-Gal4,Ds-GFP/+;UAS-hopRNAi/+; (G) btl-Gal4,Ds-GFP/+;UAS-stat92ERNAi/+; (I) btl-Gal4,Fj-GFP/+; (J) btl-Gal4,Fj-GFP/+;UAS-domeRNAi/+; (K) btl-Gal4,Fj-GFP/+;UAS-hopRNAi/+; (L) btl-Gal4,Fj-GFP/+;UAS-stat92ERNAi/+; (N) btl-Gal4/+;Ft-GFP/+; (O) btl-Gal4/+;Ft-GFP/UAS-domeRNAi; (P) btl-Gal4/+;Ft-GFP/UAS-hopRNAi; (Q) btl-Gal4/+;Ft-GFP/UAS-stat92ERNAi.

Added in the revised version: To further validate these putative Stat92E targets and investigate their dependence on JAK/STAT signaling, we analyzed their expression from several fosmid transgenes which have a GFP tag fused to ds, fj or ft and express at endogenous levels. Ds and Fj were abundant in the progenitor cells, but were vastly reduced upon depletion of dome, hop or stat92E, suggesting that they are regulated by JAK/STAT pathway (Figure 4D-M). Furthermore, it is reported that the function of Ft is influenced by cell-autonomous increase of Ds level and its protein level is enhanced by Ds reduction (Ambegaonkar et al., 2012; Matakatsu and Blair, 2004), which is also evidenced by our analysis using dsRNAi and UAS-ds in the tracheal progenitors (Figure 4-figure supplement 2A-D). In accordance with this notion, the level of Ft, as assayed by the Ft-GFP reporter, was elevated by the reduction of JAK/STAT signaling (Figure 4N-R). We also analyzed GFP-tagged fosmid transgenes of fat2, crb and yurt and found that they were discernably reduced upon impairment of JAK/STAT signaling, suggesting that they are also regulated by JAK/STAT pathway (Figure 4-figure supplement 1F-T). Additionally, the transcription of ds, fj, ft, fat2, crb, and yurt was compromised by expression of stat92ERNAi (Figure 4-figure supplement 1U). In sum, these results suggest that JAK/STAT promotes components involved in the establishment of polarity in tracheal cells.

The roles of JAK/STAT targets in the disciplined migration

To evaluate the functional roles of the polarity proteins in tracheal progenitor migration, we perturbed their expression in the tracheal progenitors by expressing RNAi against ds, ft or fj, which were identified by ChIP-seq as the targets of JAK/STAT. In these flies, tracheal progenitors exhibited bidirectional movement, which is reminiscent of the impairment of JAK/STAT signaling (Figure 5A-E and Video 5). Similar observations were obtained by over-expression of ft or ds in the trachea (Figure 5-figure supplement 1A-D), consistent with previous reports that both loss- and gain-of-function of PCP components disrupt the planar cell polarity (Adler et al., 2000; Tree et al., 2002; Vinson and Adler, 1987). The disciplined migration of tracheal progenitors was also impaired by the expression of fat2RNAi, crbRNAi, yurtRNAi or scbRNAi (Figure 5-figure supplement 1E-I), but was not affected by perturbation of molecules involved in cell adhesion such as Enabled (Ena), Fak, E-cadherin and Robo2 (Figure 5-figure supplement 1J).

Disciplined migration requires planar cell polarity system.

Earlier version: (A-E) Migration of tracheal progenitors. The migration of progenitors in control (A-A’’), dsRNAi (B-B’’), ftRNAi (C-C’’) and fjRNAi (D-D’’) flies. (E) Bar graph plots the migration distance of anterior movement. Error bars represent SEM, n=4. (F-I) Level of Ft in tracheal progenitors of control (F, G) and stat92ERNAi (H-J) flies. The images show progenitors at 1hr APF (H, H’) and 2hr APF (I, I’). Ft-GFP (green) (F, H, I), phalloidin (magenta), Hoechst (blue), and merged images (F’, H’, I’). (G, J) Profile plots showing the level of Ft-GFP in control (G) and stat92ERNAi (J) flies, n=3. The levels of Ft were measured along the dotted lines in F’ or I’. Anterior (A) and posterior (P). (K) Representative traces plot the migration distance relative to the origin, n=12. The x-axis represents the number of captured images. Individual frame is captured every 5 minutes. (K’) Rose plot depicting the direction of cell movement. (L) Representative traces showing the movement of individual fj-KO cells relative to their origin, n=11. The x-axis represents the number of captured images. Individual frame is captured every 5 minutes. (L’) Rose plot depicting the movement direction of fj-KO cells. (M) Scatter plots represent the ratio (d/D) of straight-line length displacement (d) relative to the length of the migration track (D) of individual cell. Error bars represent SEM. N.S. indicates not significant. Scale bar: 200 μm (A-D’’), 50 μm (F, F’), 100 μm (H-I’). Genotypes: (A-A’’) btl-Gal4/+;P[B123]-RFP-moe/+; (B-B’’) btl-Gal4/UAS-dsRNAi;P[B123]-RFP-moe/+; (C-C’’) btl-Gal4/+;P[B123]-RFP-moe/UAS-ftRNAi; (D-D’’) btl-Gal4/UAS-fjRNAi;P[B123]-RFP-moe/+; (F, F’) btl-Gal4/+;Ft-GFP/+; (H, H’, I, I’) btl-Gal4/+;Ft-GFP/UAS-stat92ERNAi.

Revised version: (A-E) Migration of tracheal progenitors. The migration of progenitors in control (A-A’’), dsRNAi (B-B’’), ftRNAi (C-C’’) and fjRNAi (D-D’’) flies. (E) Bar graph plots themigration distance of anterior movement. Error bars represent SEM, n=4. (F-I) Level of Ft in tracheal progenitors of control (F, G) and stat92ERNAi (H-J) flies. The images show progenitors at 1hr APF (H, H’) and 2hr APF (I, I’). Ft-GFP (green) (F, H, I), phalloidin (magenta), Hoechst (blue), and merged images (F’, H’, I’). (G, J) Profile plots showing the level of Ft-GFP in control (G) and stat92ERNAi (J) flies, n=5. ANOVA test: p < 0.0001. The levels of Ft were measured along the dotted lines in F’ or I’. Anterior (A) and posterior (P). (K) Representative traces plot the migration distance relative to the origin, n=12. The x-axis represents the number of captured images. Individual frame is captured every 5 minutes. (K’) Rose plot depicting the direction of cell movement. (L) Representative traces showing the movement of individual fj-KO cells relative to their origin, n=11. The x-axis represents the number of captured images. Individual frame is captured every 5 minutes. (L’) Rose plot depicting the movement direction of fj-KO cells. (M) Scatter plots represent the ratio (d/D) of straight-line length displacement (d) relative to the length of the migration track (D) of individual cell. Error bars represent SEM. N.S. indicates not significant. Scale bar: 200 μm (A-D’’), 50 μm (F, F’), 100 μm (H-I’). Genotypes: (A-A’’) btl-Gal4/+;P[B123]-RFP-moe/+; (B-B’’) btl-Gal4/UAS-dsRNAi;P[B123]-RFP-moe/+; (C-C’’) btl-Gal4/+;P[B123]-RFP-moe/UAS-ftRNAi; (D-D’’) btl-Gal4/UAS-fjRNAi;P[B123]-RFP-moe/+; (F, F’) btl-Gal4/+;Ft-GFP/+; (H, H’, I, I’) btl-Gal4/+;Ft-GFP/UAS-stat92ERNAi.

Earlier version: Migratory cells generate protrusions at the leading edge to initiate movement (Cetera et al., 2014). The normal posteriorly migrating tracheal progenitors extend protrusions towards the migratory directions (Figure 5-figure supplement 2A), but in the bidirectionally moving progenitors in which upd2 in fat body was perturbed, extensive filopodia were projected from both the anterior and posterior fronts (Figure 5-figure supplement 2B, C), indicating that the aberrantly anteriorly moving progenitors may adopt the identity as those moving posteriorly. Bidirectionally migrating progenitors induced by perturbation of JAK/STAT signaling did not alter the expression of the tracheal inducer, branchless (bnl) (Figure 5-figure supplement 2D, E). Further analysis revealed that the progenitors exhibited elevated levels of Ft at the leading edge where they attached to dorsal trunk (DT) (Figure 5F-G). Accordingly, progenitors that underwent bidirectional movement exhibited pronounced abundance of Ft at both the anterior and posterior frontal edges (Figure 5H-J). To further evaluate the functional roles of Ft-Ds-Fj module in disciplined migration, we utilized the high-mobility carcinoma cells, SKOV-3, and found that perturbation of Fj that phosphorylates the extracellular cadherin domains of both Ft and Ds and modifies their heterophilic binding (Thomas and Strutt, 2012), Ft or Ds concurrently displayed compromised directionality and reduced consistency of movement in a two-dimensional culture (Figure 5K-M and Figure 5-figure supplement 3). Together with the results in previous sections, these observations suggest that the activated tracheal progenitors establish a disciplined migration through the asymmetrical distribution of polarity proteins which is directed by an Upd2-JAK/STAT signaling stemming from the remote organ of fat body.

Revised version: Migratory cells generate protrusions at the leading edge to initiate movement (Cetera et al., 2014). The normal posteriorly migrating tracheal progenitors extend protrusions towards the migratory directions (Figure 5-figure supplement 2A), but in the bidirectionally moving progenitors in which upd2 in fat body was perturbed, extensive filopodia were projected from both the anterior and posterior fronts (Figure 5-figure supplement 2B, C), indicating that the aberrantly anteriorly moving progenitors may adopt the identity as those moving posteriorly. Bidirectionally migrating progenitors induced by perturbation of JAK/STAT signaling did not alter the expression of the tracheal inducer, branchless (bnl) (Figure 5-figure supplement 2D, E). Further analysis revealed that the progenitors exhibited elevated levels of Ft at the leading edge where they attached to dorsal trunk (DT) (Figure 5F-G).

Added in the revised version: Accordingly, progenitors that underwent bidirectional movement exhibited pronounced abundance of Ft at both the anterior and posterior frontal edges (Figure 5H-J). To further evaluate the functional roles of Ft-Ds-Fj module in disciplined migration, we utilized the high-mobility carcinoma cells, SKOV-3, and found that perturbation of Fj that phosphorylates the extracellular cadherin domains of both Ft and Ds and modifies their heterophilic binding (Thomas and Strutt, 2012), Ft or Ds concurrently displayed compromised directionality and reduced consistency of movement in a two-dimensional culture (Figure 5K-M and Figure 5-figure supplement 3). Together with the results in previous sections, these observations suggest that the activated tracheal progenitors establish a disciplined migration through the asymmetrical distribution of polarity proteins which is directed by an Upd2-JAK/STAT signaling stemming from the remote organ of fat body.

Upd2 in the fat body-produced vesicle

Besides the JAK/STAT signaling, another functional class enriched for vesicle-mediated transport was prominent from our surface proteome analysis of the trachea (Figure 2H). A series of components that function in vesicle trafficking were identified. It has been reported that IL-6 cytokines tend to be encapsulated in secretory vesicles (Kandere-Grzybowska et al., 2003; Verboogen et al., 2018). To visualize Upd2 production and investigate its transportation kinetics, an upd2-mCherry transgene was developed and expressed under the control of lsp2-Gal4, which enabled tracking the dynamics of Upd2 in fat body (Figure 6A). In agreement with Upd2 being transported through vesicles, administration of L3 larvae with Brefeldin A (BFA), which pharmacologically inhibits vesicle formation and transport, sequestered Upd2 proteins in fat body (Figure 6B, C). To track the destination of the Upd2-containing vesicles, we examined mCherry signals in adjacent tissues and detected considerable amount of Upd2 puncta in the tracheal progenitors (Figure 6D). BFA treatment reduced Upd2-mCherry puncta in the tracheal progenitors, suggesting that tracheal progenitors receive vesicular Upd2 from the fat body (Figure 6E, F). Perturbation of Grasp65, a Golgi reassembly stacking protein previously implicated in Upd2 secretion (Rajan et al., 2017), also led to sequestration of Upd2-containing vesicles in fat body (Figure 6G, H, J). The vesicle formation, function and extracellular movement are dependent on the tetraspanin superfamily proteins (Andreu and Yanez-Mo, 2014). We surveyed all the tetraspanin orthologs in fly for potential roles in Upd2 vesicle formation and transport. When expressing lbmRNAi in fat body, Upd2-containing vesicles were vastly increased (Figure 6I, J). Meanwhile, perturbation of vesicle secretion or transport by expressing grasp65RNAi or lbmRNAi in fat body eliminated the presence of fat body-origin Upd2 in the trachea, suggesting that fat body-produced Upd2-containing vesicles function cell non-autonomously and contribute to other tissues/organs (Figure 6K-N). It should be noted that knockdown of upd2 in the trachea did not alter the discipline of tracheal progenitor migration (Figure 2-figure supplement 1K-M). Collectively, these results suggest that fat body-produced Upd2 undergoes vesicle-mediated trafficking.

The production and transport of Upd2 from fat body.

(A, B) The number of Upd2-mCherry-containing vesicles in fat body of control (DMSO-fed, A) and BFA-treated L3 larvae (B). Larger view in lower magnification is provided in Figure 6-figure supplement 1. (C) Bar graph plots the number of Upd2-mCherry-containing vesicles in fat body. Error bars represent SEM, n=6. (D, E) The confocal image showing the number of Upd2-containing vesicles in progenitors of DMSO-fed control (D) and BFA-treated flies (E). (F) Bar graph plots the number of Upd2-mCherry-containing vesicles in progenitors. Error bars represent SEM, n=4. (G) The number of Upd2-mCherry containing vesicles (red) in fat body. (H, I) Upd2 accumulation in fat body was increased in the presence of grasp65RNAi (H) and lbmRNAi (I). (J) Bar graph plots the number of Upd2-mCherry-containing vesicles. Error bars represent SEM, n=4. (K-M) The Upd2 vesicles (red) in tracheal progenitors (DAPI) in control (K), grasp65RNAi (L) and lbmRNAi (M) flies. Dashed lines outline tracheal progenitors. (N) Bar graph plots the number of Upd2-mCherry-containing vesicles in progenitors. Error bars represent SEM, n=4. Scale bar: 20 μm (A, B, G-I), 50 μm (D, E, K-M). Genotypes: (A-G, K) UAS-upd2-mCherry/+;lsp2-Gal4/+; (H, L) UAS-upd2-mCherry/UAS-grasp65RNAi;lsp2-Gal4/+; (I, M) UAS-upd2-mCherry/+;lsp2-Gal4/UAS-lbmRNAi.

The vesicular transport in JAK/STAT signaling

The results in previous section suggest that the ligand of JAK/STAT signaling is transported in a manner that depends on vesicle trafficking. To validate the role of fat body-produced vesicles in inter-organ signaling, we used genetic and pharmacological tools to perturb different processes of vesicle trafficking in fat body and monitored JAK/STAT signaling in the tracheal progenitors. Expression of grasp65RNAi in fat body reduced the activity of JAK/STAT signaling in the trachea, as assessed by the Stat92E-GFP reporter (Figure 7A, B). Similarly, RNAi targeting expression of lbm in fat body vanished JAK/STAT signal transduction in the trachea (Figure 7C). Rab GTPases coordinate vesicle trafficking and production (Stenmark, 2009) and have been shown to play pivotal roles in the regulation of intracellular trafficking of FGFR and EGFR (Letizia et al., 2023; Olivares-Castineira and Llimargas, 2017), and were identified in the surface proteome analysis. Consistently, knockdown of rab5 or rab7 in fat body reduced the activity of JAK/STAT signaling in the progenitor cells (Figure 7D-F). Corroborating the genetic manipulations, BFA treatment that impeded vesicular transport also resulted in impairment of JAK/STAT signaling in trachea (Figure 7G-I). Taking advantage of the aforementioned DIPF reporter to assess the response of receiving cells to ligands, we found that the fluorescent signal of DIPF was compromised upon the presence of BFA, but was unaffected by inhibitors that target the downstream JAK protein (Figure 7J-M), suggesting that signaling ligands are less abundant in the recipient progenitor cells and that the vesicle-mediated transport of ligands is essential for JAK/STAT signaling.

The dependence of tracheal progenitors on vesicle trafficking.

(A-E) The expression of Stat92E-GFP in tracheal progenitors of control (A), grasp65RNAi (B), lbmRNAi (C), rab5RNAi (D) and rab7RNAi (E) flies. The progenitors are outlined by dashed lines. (F) Bar graph plots the relative expression of Stat92E-GFP. Error bars represent SEM, n=7. (G, H) The expression of Stat92E-GFP in tracheal progenitors in DMSO-fed (G) and BFA-treated (H) flies. (I) Bar graph plots the relative expression of Stat92E-GFP. Error bars represent SEM, n=5. (J) The signal of DIPF reporter in tracheal progenitors. (K) The effects of Tofacinib (JAK inhibitor) on DIPF reporter in progenitors. (L) The effects of Brefeldin A on DIPF reporter in progenitors. Dashed lines outline tracheal progenitors. (M) Bar graphs showing the signal of DIPF reporter. Error bars represent SEM, n=4. (N-O’’) Migration of tracheal progenitors in DMSO-fed flies (N-N’’) and BFA-treated flies (O-O’’). (P) Bar graph plots migration distance of anterior movement. Error bars represent SEM, n≥3. (Q-U’’) Migration of tracheal progenitors at 0hr APF (Q), 1hr APF (Q’), 3hr APF (Q’’). The confocal images showing the tracheal progenitors in control (Q-Q’’), grasp65RNAi (R-R’’), lbmRNAi (S-S’’), rab5RNAi (T-T’’) and rab7RNAi (U-U’’) flies. (V) Bar graph plots the migration distance of anterior movement. Error bars represent SEM, n=4. Scale bar: 50 μm (A-E, G, H, J-L), 200 μm (N-O’’, Q-U’’). Genotypes: (A, G, H) lsp2-Gal4,Stat92E-GFP/+; (B) UAS-grasp65RNAi/+;lsp2-Gal4,Stat92E-GFP/+; (C) lsp2-Gal4,Stat92E-GFP/UAS-lbmRNAi; (D) UAS-rab5RNAi/+;lsp2-Gal4,Stat92E-GFP/+; (E) UAS-rab7RNAi/+;lsp2-Gal4,Stat92E-GFP/+; (J-L) btl-Gal4/UAS-DIPF; (N-Q’’) lsp2-Gal4,P[B123]-RFP-moe/+; (R-R’’) lsp2-Gal4,P[B123]-RFP-moe/UAS-grasp65RNAi; (S-S’’) lsp2-Gal4,P[B123]-RFP-moe/+;UAS-lbmRNAi/+; (T-T’’) UAS-rab5RNAi/+;lsp2-Gal4,P[B123]-RFP-moe/+; (U-U’’) UAS-rab7RNAi/+;lsp2-Gal4,P[B123]-RFP-moe/+.

Phenotypically, the tracheal progenitors exhibited bidirectional migration in BFA-treated flies, which phenocopies JAK/STAT loss-of-function (Figure 7N-P). In concord with this observation, depletion of grasp65 or lbm also led to bidirectional movement (Figure 7Q-S’’, and Video 6). Similar observations were made in tracheal progenitors with either rab5 or rab7 knockdown (Figure 7T-V), whereas perturbation of neither rab2 nor rab3 affected the disciplined progenitor migration (Figure 7-figure supplement 1A-C). Taken together, these results suggest that JAK/STAT signaling in the trachea is dependent on the vesicle-mediated transport of its ligands from fat body.

The interaction between Upd2 and endocytic machinery

Our results described thus far suggest that Upd2 emanating from fat body signals to JAK/STAT signaling in the trachea. To further explore the molecular basis underlying the vesicular transport of Upd2, we monitored Rab5-GFP and Rab7-GFP in fat body, which mark early and late endosomes, respectively (Vonderheit and Helenius, 2005). The fat body-produced Upd2 appeared vesicular (Figure 8A) and both Rab5 and Rab7 were found adjacent to the Upd2-harboring vesicles, suggesting that both Rab GTPases function in the transport of Upd2 (Figure 8A’-B’’). In contrast, Rab3 exhibited non-overlapping distribution with Upd2 (Figure 7-figure supplement 1E-E’’). Furthermore, we observed that Grasp65 was in close proximity to Upd2-containing vesicles, indicating its integral roles in these vesicles (Figure 8C-C’’). In addition, Upd2 was observed to colocalize with the tetraspanin, Lbm (Figure 8D-D’’). At higher resolution, Upd2 and Lbm showed close association in a supramolecular configuration (Figure 8D’’’and 8D’’’’), corroborating its role in the transport of Upd2. To determine if Upd2 interacts with the coordinators of vesicle trafficking, we employed the Duolink in situ proximity ligation assay (PLA) which revealed strong interactions between Upd2 and Rabs, such as Rab5 and Rab7 (Figure 8E-H), as well as Lbm (Fig 8I, J). The interaction was further validated by the revelation that Upd2 co-immunoprecipitated with Rab5 and Rab7 (Figure 8K, L). The presence of Upd2 in Lbm-containing vesicles was also evidenced in S2 cells (Figure 8-figure supplement 1A-C’’) and co-IP experiment showed that Lbm physically associated with Upd2 in both fat body and S2 cells (Figure 8M and Figure 8-figure supplement 1F). To further understand the biogenesis of Lbm-containing vesicles that transport Upd2, we conducted electron microscopic analysis of the Lbm-containing vesicles through the expression of an HRP-fused Lbm in the fat body (Figure 8-figure supplement 1D). The interaction between Lbm and Upd2 as simulated by Alphafold2 supported their direct association (Figure 8-figure supplement 1E). Then, we generated an Lbm chimera tagged with a pH-sensitive GFP variant, pHluorin (Yoshihara et al., 2005). PHluorin fluorescence is squelched at the low pH domain such as in intravesicular compartments, but becomes detectable when exposed to the extracellular environment, thus enabling detection of exocytosis and endocytosis. Fat body expressing Lbm-pHluorin produced GFP puncta at the plasma membrane (Figure 8N), and the GFP signal was also detected in the trachea, suggesting the reception and internalization of Lbm-containing vesicles by tracheal cells (Figure 8O). However, the GFP fluorescence in both fat body and trachea was dramatically decreased by BFA treatment, suggesting that Lbm-containing vesicles are diminished (Figure 8-figure supplement 1G-J). Accordingly, the signals of Lbm-pHluorin in both fat body and responding tracheoblasts were apparently compromised when Rab5 or Rab7 was perturbed (Figure 8P-S), suggesting that the biogenesis and production of Lbm-containing vesicles depend on Rab-mediated vesicle trafficking. Taken together, these results suggest that fat body-derived Upd2 interacts with Rab-mediated endocytic trafficking system to control the disciplined movement of tracheal progenitors.

The roles of endocytic trafficking system in the transport of Upd2.

(A-D’’’) The confocal images showing the colocalization between Upd2 (red) (A, B, C, D) and Rab5 (GFP) (A’), Rab7 (GFP) (B’), Grasp65 (GFP) (C’) or Lbm (GFP). (A’’, B’’, C’’, D’’) Merged images. (D’’’) 3D high-magnification view of the boxed inset in D’’. (D’’’’) The Pearson’s correlation coefficient depicting colocalization between Lbm and Upd2 in fat body cells. (E-J) The PLA (Proximity Ligation Assay) assay showing the interaction between Upd2 and Rab5 (E, F), Rab7 (G, H) or Lbm (I, J). (K-M) Co-immunoprecipitation assay showing physical interaction between Upd2 and Rab5 (K), Rab7 (L) or Lbm (M) in larval fat body. (N-T) The expression of Lbm-pHluorin in larval fat body and progenitors of control (N, O), rab5RNAi (P, Q) and rab7RNAi (R, S) flies. Dashed lines outline tracheal progenitors. Arrowheads point to Lbm-pHluorin puncta. DAPI signal indicates nuclei. (T) Schematic diagram depicting Upd2-operated disciplined migration of tracheal progenitors. Scale bar: 5 μm (A-D’’, E-J), 10 μm (N, P, R), 20 μm (O, Q, S). Genotypes: (A-A’’, E, F) UAS-upd2-mCherry/+;lsp2-Gal4/UAS-GFP-rab5; (B-B’’, G, H) UAS-upd2-mCherry/+;lsp2-Gal4/UAS-GFP-rab7; (C-C’’) UAS-upd2-mCherry/UAS-grasp65-GFP;lsp2-Gal4/+; (D-D’’’’, I-L) UAS-upd2-mCherry/UAS-lbm-GFP;lsp2-Gal4/+; (N, O) UAS-lbm-pHluorin/+;lsp2-Gal4/+; (P, Q) UAS-lbm-pHluorin/UAS-rab5RNAi;lsp2-Gal4/+; (R, S) UAS-lbm-pHluorin/UAS-rab7RNAi;lsp2-Gal4/+.

Discussion

Resident stem cells and progenitors are mobilized to regenerate damaged or degenerated tissue. Despite the large distance between the niche where stem cells interact with their microenvironments and the destination for reconstruction, their commitment to a stereotyped track implicates sophisticated mechanism that controls disciplined migration as stem cells are activated and move out of the niche. While primary inducers expressed by damaged tissues are in play to coordinate the newly generated architecture with the degenerated counterpart (Chen and Krasnow, 2014), the present study elucidates an integral role of Upd2-JAK-STAT pathway in regulating the expression of polarity-related genes and maintaining the disciplined migration of tracheal progenitors (Figure 8T). The transport of Upd2 from fat body to trachea suggests intensive inter-organ communication during the migration of tracheal progenitors.

Several possibilities could account for the JAK/STAT-dependent polarity. The signaling components of JAK/STAT pathway could exhibit polarized localization (Sotillos et al., 2008). Alternatively, this signaling may activate genes controlling cell polarity and adhesion (Mallart et al., 2024; Tsurumi et al., 2011). Functional interplay between JAK/STAT signaling and cell polarity has been observed in various contexts (Chatterjee et al., 2023; Zeidler et al., 1999). Our results support a role of JAK/STAT signaling in promoting expression of genes with established roles in planar polarity, which may hallmark the route for the migration of the tracheal progenitor cells. Epithelial cells exhibit two aspects of polarity: apical-basal polarity and planar cell polarity (PCP). The latter refers to the collective alignment of cell polarity within the plane of an epithelial sheet (Zallen, 2007). Molecularly, PCP is generated by the asymmetry of a group of proteins (PCP proteins) that mediate communication between neighboring cells (Barlan et al., 2017a; Matis and Axelrod, 2013; Williams et al., 2022). In Drosophila, the components of PCP are considered to group functionally into two core modules. The seven-pass transmembrane protein Frizzled (Fz), the cytosolic proteins Dishevelled (Dsh), Diego (Dgo), the four-pass transmembrane protein Strabismus (Stbm, also known as Van Gogh (Vang)) and the cytosolic protein Prickle (Pk) belong to the first module. The second module consists of Fat (Ft; also known as cadherin-related tumor suppressor), Dachsous (Ds), Four-jointed (Fj) and Atro (a transcription repressor) (Peng and Axelrod, 2012). Intensive functional interplay occurs between these two modules (Ayukawa et al., 2014). Aberrant activity of the core PCP proteins leads to misoriented hairs and complex swirling patterns (Ma et al., 2003; Wong and Adler, 1993). In addition to arrangement of epithelial appendages, PCP pathway is also required for collective and directed cell movements (Munoz-Soriano et al., 2012). The migratory cell cohort is polarized into ‘pioneer’ cells that lead the trailing followers (Vitorino and Meyer, 2008).

Our data indicate that expression of Ds, Fj, Fz, Stan and Fat2, core components or regulators of PCP, depends on JAK/STAT pathway. The interaction between atypical cadherin Fat (Ft) and its ligand, Dachsous (Ds) directs core protein asymmetry (Strutt and Strutt, 2021; Yang et al., 2002). Phenotypically, aberrancy of PCP protein abundance, either excessive core protein or deficit in expression gradients, gives rise to similar morphological abnormality (Adler et al., 2000; Casal et al., 2002; Taylor et al., 1998; Tree et al., 2002). Consistent with this notion, gain-of-function of Fj or Ds phenocopies that of perturbation of PCP proteins. JAK/STAT signaling promotes the expression of Ds, but reduces Ft expression (Figure 4). Thereby, perturbation of JAK/STAT signaling disrupts the Ds-Ft system.

A precedent for JAK/STAT signaling in directional cell movement is border cell migration from anterior to posterior compartment during Drosophila oogenesis. Migration of the border cells is guided by a gradient of PDGF and VEGF chemokines (Duchek et al., 2001). Loss of either hop (encoding JAK) or stat in the border cells impinges their recruitment into the cluster and subsequent migration (Silver et al., 2005). Our results suggest that JAK/STAT signaling does not serve as a guidance cue for tracheal progenitors, but rather directs the directionality of cell movement. The downstream PCP components may contribute to either polarity of progenitors or cell-cell interactions between the progenitors and tracheal cells that they track along. It remains unknown how individual progenitor cells perceive directional information and convert it into group choreography.

We identified the fly fat body as the major source for the JAK/STAT signaling ligand, Upd2 production. Fat body is functionally equivalent to the mammalian liver which stores proteins, lipids and sugars and functions as an energy reservoir (Li et al., 2019). It supplies proteins and/or hormones that are utilized by other organs, and thereby serves as an interchange center to disperse systemic hormonal and nutritional signals. For instance, it generates collagen IV to decorate imaginal discs and produces xanthine dehydrogenase for eye pigmentation (Pastor-Pareja and Xu, 2011; Reaume et al., 1989). The transport between fat body and trachea has been reported on a secreted chitin deacetylase, Serpentine (Serp), which is expressed by fat body and contributes to tracheal morphogenesis (Dong et al., 2014). Our results reveal that fat body also signals to regulate the disciplined migration of tracheal progenitors through the dispersion of Upd2 cytokines. These studies collectively suggest that fat body orchestrates systemic tissue growth and patterning and that metabolic regulation is critical for adult stem cells.

Proteins that are locally produced can execute systemic function in distant organs. A possible route of transport is through the hemolymph or bloodstream and taken up by the target tissues. The signaling proteins such as cytokines can be packaged in extracellular vesicles with various dimensions (Buzas, 2023; Javeed et al., 2021). A precedent of vesicular transport of signaling molecules is reported in migrasomes whose diameter exceeds 500 nm (Jiang et al., 2019).

These extracellular vesicles mediate cell-to-cell communication (Colombo et al., 2014), perhaps at a distance (Hood et al., 2011) and even traverse between organs (Corrigan et al., 2014). The vesicular Upd2 is able to signal at recipient cells, suggesting that the activity of Upd2 is preserved in the vesicle, and it is released upon vesicle fusion. Compared with conventional extracellular vesicles such as exosomes, the Upd2-containing vesicles possess larger dimension. Its production and trafficking depend on GRASP-mediated unconventional secretion and interaction with Lbm. Lbm belongs to the tetraspanin protein family that contains four transmembrane domains. The mammalian homologs of tetraspanins, CD9, CD63, CD81 or CD37 are principal constituents of extracellular vesicles. The Lbm-containing vesicles are regulated by GRASP-mediated secretion and are sensitive to pharmacological inhibition of EV transport.

It has been proposed that tetraspanins facilitate regeneration and wound healing in cultural cells and single-cell plasma membrane. Tetraspanin-enriched macrodomains are assembled into a ring-like structure (Huang et al., 2022), which is recruited to large membrane wounds and promotes membrane repair (Wang et al., 2022). The present study adds another dimension to the roles of tetraspanin proteins in tissue regeneration which can be ascribed to transport of signaling proteins and modulation of stem cells.

Materials and methods

Key Resources Table

Fly lines and husbandry

Earlier version: All flies were reared on normal cornmeal and agar medium at 25L unless noted. UAS-upd2RNAi (THU1331, THU1288), UAS-domeRNAi (THU0574), UAS-hopRNAi (THU5762), UAS-stat92ERNAi (THU0573), UAS-lbmRNAi (THU2602), UAS-grasp65RNAi (TH04282.N), UAS-rab5RNAi (TH02192.N), UAS-rab7RNAi (TH02539.N), UAS-fjRNAi (THU201500988.S), UAS-fat2RNAi (THU4120), UAS-yurtRNAi (THU1740), UAS-crbRNAi (THU2783) and UAS-scbRNAi (THU3905) were ordered from Tsinghua Stock Center. Stat92E-GFP (BSDC#38670), UAS-grasp65-GFP (BSDC#8507), Ds::GFP (BSDC#59425) were from Bloomington Drosophila Stock Center. Yurt::GFP (v318067), Crb::GFP (v318384), Ft::GFP (v318477), Fj::GFP (v318457), UAS-dsRNAi (v36219, THU2846), and UAS-ftRNAi (v9396) were obtained from VDRC. UAS-ft was kindly provided by Dr. Xianjue Ma, UAS-ds was kindly provided by Dr. Xing Wang and UAS-fat2-GFP was from Dr. Shunfan Wu. UAS-GFP-rab5 and UAS-GFP-rab7 were kindly provided by Dr. Xiaohang Yang.

Revised version: All flies were reared on normal cornmeal and agar medium at 25℃ unless noted. UAS-upd2RNAi (THU1331, THU1288), UAS-domeRNAi (THU0574), UAS-hopRNAi (THU5762), UAS-stat92ERNAi (THU0573), UAS-lbmRNAi (THU2602), UAS-grasp65RNAi (TH04282.N), UAS-rab5RNAi (TH02192.N), UAS-rab7RNAi (TH02539.N), UAS-fjRNAi (THU201500988.S), UAS-fat2RNAi (THU4120), UAS-yurtRNAi (THU1740), UAS-crbRNAi (THU2783) and UAS-scbRNAi (THU3905) were ordered from Tsinghua Stock Center. Stat92E-GFP (BSDC#38670), UAS-grasp65-GFP (BSDC#8507), Ds::GFP (BSDC#59425) were from Bloomington Drosophila Stock Center. Yurt::GFP (v318067), Crb::GFP (v318384), Ft::GFP (v318477), Fj::GFP (v318457), UAS-dsRNAi (v36219, THU2846), and UAS-ftRNAi (v9396) were obtained from VDRC. UAS-ft was kindly provided by Dr. Xianjue Ma, UAS-ds was kindly provided by Dr. Xing Wang and UAS-fat2-GFP was from Dr. Shunfan Wu. UAS-GFP-rab5 and UAS-GFP-rab7 were kindly provided by Dr. Xiaohang Yang.

Plasmid construction and transgenic flies

To generate UAS-upd2-mCherry and UAS-lbm-GFP, UAS-lbm-HRP transgenic flies, the coding sequence of upd2 or lbm was PCR amplified from a fly cDNA library and cloned into a pUAST vector with C-terminal mCherry, GFP or HRP.

The DIPF reporter was generated by first fusing cpYFP and Drosophila FKBP12 and subsequently ligating to the dome cDNA via a GTG linker. The above product was then cloned into a pUAST vector and verified by DNA sequencing, and injected into y[1] M{vasint.Dm}ZH-2A w[*]; P{CaryP}attP2 recipient flies or w1118 following standard Drosophila transformation injection procedures (Core Facility of Drosophila Resource and Technology, SIBCB, CAS).

Cell culture and transfection

S2 cells were grown in Schneider Drosophila Medium (Gibco, #21720024) supplemented with 10% (v/v) Fetal Bovine Serum (FBS, Gibco, #10099141C) and 1% (v/v) Penicillin Streptomycin (Pen/Strep, Life Technologies) at 28°C with 0.2% CO2. Transfection was conducted with 5 μg plasmids (act-GAL4, UAS-lbm-GFP, UAS-upd2-mCherry) using Effectene Transfection Reagent (QIAGEN, #301425).

SKOV3 cells were cultured in DMEM medium (CR#12800) containing 10% FBS (Gibco, #10099141C) in an incubator with 5% CO2 at 37 °C. Cells with ∼80% confluency were infected with lentivirus loaded with siRNAs for gene knockdown or gRNAs for knockout, and the medium was replaced 24 hours post-infection. 72 hours after lentivirus infection, bright-field images were taken every 5 minutes for 12 hours using a confocal microscope.

Quantitative reverse transcription PCR

Larval or pupal trachea were dissected in cold PBS, and then transferred to RNA extraction reagent (AG21101). Next, reverse transcription was performed using qPCR RT Mix with gDNA Remover reagent (AG11706). qPCR was performed using the Universal SYBR Select Master Mix (AG11701) with a Bio-Rad system. The foldchange of target gene expression was normalized to that of α-tubulin. The primers are listed in Key Resources Table.

Western Blotting and Co-immunoprecipitation

Total protein was extracted from cells or tissues by RIPA buffer supplemented with a protease inhibitor cocktail (Merck, #11836170001) and phenylmethanesulfonyl fluoride (Beyotime, #ST507), separated by 10% SDS-PAGE gels and transferred to PVDF membrane (Millipore, #IPVH00010). Blots were detected with an ECL Western Blotting detection system (Bio-Rad). For co-immunoprecipitation, lysates of larval fat body or transfected S2 cells were incubated overnight at 4°C with protein A Magnetic beads (Thermo Scientific, #2736141) pre-coated with GFP antibody (Invitrogen, #A11122). Immunoprecipitates were eluted in SDS-containing loading buffer for subsequent immunoblotting analysis. Antibodies for immunoblotting include: α-tubulin (Baoke, #BK7010), α-HRP-conjugated Streptavidin (Proteintech, #SA00001-0), α-GFP (Abclonal, #AE012, 1:1000); α-mCherry (Abclonal, #AE002, 1:1000).

Proximity Ligation Assay

Proximity ligation assay (PLA) was carried out with Duolink® In Situ Detection Reagents Far Red (Sigma Aldrich, #DUO92013) according to manufacturer’s instructions, using the probes anti-rabbit PLUS (Sigma Aldrich, #DUO92002) and anti-mouse MINUS (Sigma Aldrich, #DUO92004). Briefly, larval fat body was dissected and fixed in 4% formaldehyde. The animals not expressing upd2-mCherry served as controls. After permeabilization, the samples were incubated with primary antibodies overnight at 4°C. Then, the samples were washed with PLA buffer A, hybridized with PLA probes, ligated, and amplified. Samples were washed twice with PLA buffer B (Sigma Aldrich, #DUO82049) and fluorescence images were taken with an LSM Zeiss 900 inverted confocal laser scanner microscope.

Immunofluorescence

Trachea from white pupae (0hr APF) were dissected in PBS and fixed with 4% formaldehyde for 25 min at room temperature. After washes, trachea samples were permeabilized with 1% TritonX-100 in PBS, and then blocked in 10% goat serum. Incubation with primary antibody (GFP, 1:400; lacZ, 1:40) was performed at 4°C with gentle rotation for overnight. Then, the samples were incubated with secondary antibodies conjugated to Alexa Fluor 488 or 555 (1:200) and Phalloidin (1:50) for 2 hrs. After washing, samples were mounted in antifade mounting medium with DAPI (VECTASHIELD) and imaged under an LSM Zeiss 900 inverted confocal laser scanner microscope.

Live imaging of pupal trachea stem cells

Earlier version: White pupae (0hr APF) were briefly washed in double distilled water and mounted in halocarbon oil 700 (Sigma, #H8898). The pupae were positioned with forceps to bring a single dorsal trunk of the trachea up for optimal imaging of Tr4 and Tr5 metameres. Then, pupae were immobilized by a 22×30 mm No.1.5 high precision coverslip spaced by vacuum grease. Time-lapse images were captured by an LSM Zeiss 900 inverted confocal laser scanner microscope. For migration distance measurement, we took sequential snapshots of the moving progenitors of pupae staged at 0hr, 1hr, 2hr and 3hr APF. The migration distance was measured as the distance from the starting position (the junction of TC and DT) to the leading edge of progenitor groups. The migration velocity was calculated by v=d (micrometer)/t (min).

Revised version: White pupae (0hr APF) were briefly washed in double distilled water and mounted in halocarbon oil 700 (Sigma, #H8898). The pupae were positioned with forceps to bring a single dorsal trunk of the trachea up for optimal imaging of Tr4 and Tr5 metameres. Then, pupae were immobilized by a 22x30 mm No.1.5 high precision coverslip spaced by vacuum grease. Time-lapse images were captured by an LSM Zeiss 900 inverted confocal laser scanner microscope. For migration distance measurement, we took sequential snapshots of the moving progenitors of pupae staged at 0hr, 1hr, 2hr and 3hr APF. The migration distance was measured as the distance from the starting position (the junction of TC and DT) to the leading edge of progenitor groups. The migration velocity was calculated by v=d (micrometer)/t (min).

RNA sequencing of tracheal progenitors

Total RNA was isolated from the Tr4 and Tr5 metamere progenitors dissected from 1hr APF pupae using RNeasy Micro Kit (QIAGEN, #74004). SMART-Seq v4 Ultra low input RNA Kit (Takara, #634889) was used for first-strand and second-strand cDNA synthesis and double-stranded cDNA end repair. Double-stranded cDNAs were cleaned using AMPure XP (Beckman Coulter, #A63882). Then cDNAs were subjected to tagmentation and ligation to adaptors to generate the sequencing libraries using TruePrep DNA Library Prep Kit V2 for Illumina kit (Vazyme, #TD501). The quality and concentration of the libraries were assessed using the Agilent High Sensitivity DNA Kit and Bioanalyzer 2100 (Agilent Technologies) and submitted to 150 bp paired-end high throughput sequencing using Hiseq4000 (Illumina).

Analysis of RNA-seq data was performed using a computer system equipped with multiple processors. Clean reads were mapped to the Drosophila genome sequence using Hisat2 with default parameters. Successfully mapped reads were counted using FeatureCounts. Differential gene expression analysis was performed using the DESeq2 package. Adjusted p-value < 0.05 was used as the threshold to identify the differentially expressed genes. Gene ontology and KEGG pathway enrichment analyses for the differentially expressed genes were conducted using the Database for Annotation, Visualization and Integrated Discovery (DAVID).

Chromatin Immunoprecipitation

Third instar larval trachea from Stat92E-Flag (BDSC, #38670) were fixed in 1% formaldehyde. The fixation reaction was terminated by adding glycine (125 mM). Trachea were washed and resuspended in lysis buffer, and sonicated to generate 200-600 bp DNA fragments. Procedures of immunoprecipitation and ChIP sequencing library construction were as previously described (Li et al., 2022). Anti-FLAG M2 Magnetic Beads (Millipore, #8823) were used for enriched DNA binding to transcription factor Stat92E.

Immunoprecipitated DNA was subjected to next-generation sequencing using the Epicenter Nextera DNA Sample Preparation Kit or to real-time PCR. Library construction was performed using the High Molecular Weight tagmentation buffer (source), and tagmented DNA was linearly amplified by PCR. The libraries were then sequenced on a Novaseq according to manufacturer’s standard protocols. The sequences were processed using Fastqc and low-quality bases and adaptor contamination were trimmed by cutadapt. Filtered reads were mapped to Drosophila genome using BWA mem algorithm. Peaks were called using macs2 callpeak (Zhao et al., 2019) and plotted using pyGenomeTracks. GO analysis of biological processes was conducted by DAVID.

Cell-surface proteomics of fly trachea

Trachea from white pupae were dissected in pre-cooled Schneider Medium (Gibco) and collected in 1.5 mL low-binding tube (Axygen) containing 500 μL Schneider Medium. The samples were washed with 500 μL fresh medium and incubated with 100 μM BXXP (APEXBIO, #A8012) for 1 hr on ice with occasional pipetting. Labeling reaction was initiated by adding 1 mM (0.03%) H2O2 to the sample-containing medium and proceeded for 7 minutes at room temperature. The reaction was immediately quenched by five thorough washes with PBS containing 10 mM sodium ascorbate (Aladdin, #S105024) and 5 mM Trolox (APEXBIO, #C3183). For biochemical characterization or proteomic sample preparation, the quenching solution was drained, and the trachea in minimal residual quenching solution were quickly frozen in liquid nitrogen and stored at 80°C. LC-MS/MS analysis was performed using a Q Exactive HF-X instrument (Thermo Fisher) coupled with Easy-nLC 1200 system. The acquired MS raw data were processed using MaxQuant version 2.0.1.0 (Max Planck Institute of Biochemistry, Germany). Label-free quantification (LFQ) was set with a default parameter and iBAQ was selected.

Transmission electron microscopy

Earlier version: Fat body of third instar larvae were dissected and fixed in 0.12 M Na-cacodylate buffer (pH 7.4) containing 2.5% glutaraldehyde for 1 hr on ice. Then the samples were rinsed in 0.12 M Na-cacodylate buffer (6 × 5 min, on ice). The dissected fat body were pre-incubated with DAB (10 mg/20 ml) in 0.12 M Na-cacodylate buffer (containing 0.1% saponin) for 30 minutes with agitation in the dark. Then 30% H2O2was quickly mixed in DAB solution to a 0.03% v/v concentration and reacted for 30 min at RT. The fat body were transferred to 0.12 M Na-cacodylate buffer (6 × 5 min, RT). To increase the electron density of the HRP/DAB product, samples were transferred into 0.01% OsO4 in 0.12 M Na-cacodylate buffer (pH 7.4) for 10 min at RT, then rinsed in 0.12 M Na-cacodylate buffer (3 × 10 min, on ice). 0.1% thiocarbohydrazide in 0.12 M Na-cacodylate buffer (pH 7.4) was used for 10 min at RT, then rinsed in 0.12 M Na-cacodylate buffer (3 × 10 min, on ice). After post-fixation for 1 hour at RT with 1% OsO4 in 0.12 M Na-cacodylate buffer (pH 7.4), the samples were rinsed with MilliQ water (3 × 5 min, RT) and dehydrated in a series of 15 min with 10%, 30%, 50%, 70%, 90% and 100% (3×) ethanol. Infiltration was conducted at RT with a mixture of acetone and resin 1:1 for 1.5 h, 1:2 for 3 h and 1:3 overnight. The tissues were then dissected from the carcasses and placed in block molds filled with resin for hardening at 60°C during 48 h. 70 nm ultrathin section from the hardened blocks were cut on a Leica EM UC7 ultramicrotome using an Ultra 45° diamond knife and imaged in a thermos scientific Talos L120C electron microscope.

Revised version: Fat body of third instar larvae were dissected and fixed in 0.12 M Na-cacodylate buffer (pH 7.4) containing 2.5% glutaraldehyde for 1 hr on ice. Then the samples were rinsed in 0.12 M Na-cacodylate buffer (6 × 5 min, on ice). The dissected fat body were pre-incubated with DAB (10 mg/20 ml) in 0.12 M Na-cacodylate buffer (containing 0.1% saponin) for 30 minutes with agitation in the dark. Then 30% H2O2 was quickly mixed in DAB solution to a 0.03% v/v concentration and reacted for 30 min at RT. The fat body were transferred to 0.12 M Na-cacodylate buffer (6 × 5 min, RT). To increase the electron density of the HRP/DAB product, samples were transferred into 0.01% OsO4 in 0.12 M Na-cacodylate buffer (pH 7.4) for 10 min at RT, then rinsed in 0.12 M Na-cacodylate buffer (3 × 10 min, on ice). 0.1% thiocarbohydrazide in 0.12 M Na-cacodylate buffer (pH 7.4) was used for 10 min at RT, then rinsed in 0.12 M Na-cacodylate buffer (3 × 10 min, on ice). After post-fixation for 1 hour at RT with 1% OsO4 in 0.12 M Na-cacodylate buffer (pH 7.4), the samples were rinsed with MilliQ water (3 × 5 min, RT) and dehydrated in a series of 15 min with 10%, 30%, 50%, 70%, 90% and 100% (3×) ethanol. Infiltration was conducted at RT with a mixture of acetone and resin 1:1 for 1.5 h, 1:2 for 3 h and 1:3 overnight. The tissues were then dissected from the carcasses and placed in block molds filled with resin for hardening at 60°C during 48 h. 70 nm ultrathin section from the hardened blocks were cut on a Leica EM UC7 ultramicrotome using an Ultra 45° diamond knife and imaged in a thermos scientific Talos L120C electron microscope.

Numeration of tracheal progenitors

The pupae at indicated stages were dissected and then fixed with 4% PFA for 25 minutes. The nuclei were labeled by DAPI. The samples were imaged by an LSM Zeiss 900 inverted confocal laser scanner microscope. The number of progenitor cells was scored from image stacks.

Optic Flow Analysis

Motion Correlation

Time-lapse images of trachea progenitor cells were captured every 5 minutes over a total duration of 2 hrs using an LSM Zeiss 900 microscope. The movies and images were subjected to a three-step motion collection using ImageJ (Schneider et al., 2012) as follows:

Images were transformed to grey scale.

Automated stabilization was performed on grey images by image stabilizer.

The feature points were mapped and affine transformation was applied using big warp (Bogovic et al., 2015).

Optical flow

Earlier version: Optical flow represents the pattern of motion of pixels in a sequence of images. Between two consecutive frames I (x, y, z) and I (x, y, t + Δt), the optical flow vector v = (vX, vY) represents the motion of pixels during this time. The optical flow constraint equation is shown below: We adopted the PCA-flow algorithm (Bradski, 2000) in openCV library (cv::optflow::OpticalFlowPCAFlow) in which the sparse optical flow vectors within each small region of the image sequence are calculated before training optical flow fields via principal component analysis (PCA) (Wulff and Black, 2015). The vectors were assembled to generate a smooth vector field of optic flow using the learning linear models of flow. The PCA-flow was validated for the efficiency and robustness.

Revised version: Optical flow represents the pattern of motion of pixels in a sequence of images. Between two consecutive frames 𝐼(𝑥, 𝑦, 𝑡) and 𝐼(𝑥, 𝑦, 𝑡 + 𝛥𝑡), the optical flow vector 𝐯 = (𝑣𝑥, 𝑣𝑦) represents the motion of pixels during this time. The optical flow constraint equation is shown below: We adopted the PCA-flow algorithm (Bradski, 2000) in openCV library (cv::optflow::OpticalFlowPCAFlow) in which the sparse optical flow vectors within each small region of the image sequence are calculated before training optical flow fields via principal component analysis (PCA) (Wulff and Black, 2015). The vectors were assembled to generate a smooth vector field of optic flow using the learning linear models of flow. The PCA-flow was validated for the efficiency and robustness.

Variance in 1D axis

Earlier version: We developed a robust estimator to evaluate the variance of optic flow projection along the 1D migration axis between samples. The vector length is normalized by a frame-specific normalization factor κ to fit with scale of optic flow computed among different frames. where Q is the quantile function. The variance is computed as

Revised version: We developed a robust estimator to evaluate the variance of optic flow projection along the 1D migration axis between samples. The vector length is normalized by a frame-specific normalization factor κ to fit with scale of optic flow computed among different frames. where 𝑄 is the quantile function. The variance is computed as

Earlier version: Varc ({||v||2 > QO.2 | v ∈ vectorfield}) and scaled to 0-1 by Var(x)/(Var(x) + Var(y)).

Revised version: Var({||𝑣||2 > 𝑄0.2 | 𝑣 ∈ 𝑣𝑒𝑐𝑡𝑜𝑟𝑓𝑖𝑒𝑙𝑑}) and scaled to 0-1 by Var(𝑥)/(Var(𝑥) + Var(𝑦)).

Random Variable

The direction of optic flow in each volume was assigned to ’left’ or ’right’. Then the distribution of the binarized directions is modeled as a Bernoulli random variable X ∼ Bernoulli(p) with PMF

Binary Entropy

Earlier version: The entropy for Bernoulli random variable X ∼ Bernoulli(p) is defined as The entropy evaluates the information contained in the random variable (also called uncertainty). In this case, when p= 0.5, it reaches the maxima 1; if the variable is determinate (i.e. p=0 or 1), the entropy is zero. More directed cell migration leads to a lower entropy in optic flow since the certainty is high for the migration direction. We estimated p by computing ratio of left and right direction

Revised version: The entropy for Bernoulli random variable X ∼ Bernoulli(p) is defined as The entropy evaluates the information contained in the random variable (also called uncertainty). In this case, when 𝑝 = 0.5, it reaches the maxima 1; if the variable is determinate (i.e. p=0 or 1), the entropy is zero. More directed cell migration leads to a lower entropy in optic flow since the certainty is high for the migration direction. We estimated 𝑝 by computing ratio of left and right direction

Image and statistical analysis

Confocal z-stack images were analyzed to extract information of fluorescent intensity of stat92E-GFP and DIPF, and the number of Upd2-mCherry-containing vesicles. z slices of fluorescent intensity for Ds, Ft and Fj were measured. The number of particles for DIPF and Upd2-mCherry-containing vesicles was analyzed using ImageJ. All statistical analysis was conducted using GraphPad Prism 8.0. Mean and standard error of the mean (SEM) were shown. Unpaired t-tests with Benjamin’s correction were used to evaluate statistical significance between groups.